# Lab 1

In this lab, you will get hands-on practice with the basics of **object-oriented programming**.

### Q1: Create a `Text` class

Create a class called `Text`. 

The `__init__` function should take the following arguments (in addition to `self`):

- `title`: a `str`
- `author`: a `str`
- `year`: an `int`

```python
>>> b1 = Text("1984", "George Orwell", 1949)
```

See also the `assert` tests below.

<!-- grade_id: cell-5cad654c3635d6a3, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
class Text(object):
    def __init__(self, title, author, year):
        self.title = title
        self.author = author
        self.year = year
### END SOLUTION
```

<!-- grade_id: cell-85e53b8bdd7fd678, points: 1, grade: true, solution: false -->
```python
t1 = Text("1984", "George Orwell", 1949)
assert t1.title == "1984"
assert t1.author == "George Orwell"
assert t1.year == 1949
```

<!-- grade_id: cell-2ebcdd78699cf72e, points: 1, grade: true, solution: false -->
```python
t1 = Text("National Geographic", "Various", 1972)
assert t1.title == "National Geographic"
assert t1.author == "Various"
assert t1.year == 1972
```

### Q2: Create a `Novel` class

Now, create a new class called `Novel`, which **inherits** from `Text`. It has all the same attributes, but should add two other attributes called `genre` (a `str`) and `word_count` (an `int`).

Additionally, `Novel` should implement a `get_reading_time` method, which divides `self.word_count` by an estimated words_per_minute of `250` and rounds the answer to the nearest integer.

```python
>>> n1 = Novel("1984", "George Orwell", 1949, "dystopian", 89000)
>>> n1.get_reading_time()
356
```


<!-- grade_id: cell-bee93d5e8df6aec7, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
class Novel(Text):
    def __init__(self, title, author, year, genre, word_count):
        super().__init__(title, author, year)
        self.genre = genre
        self.word_count = word_count
        
    def get_reading_time(self):
        return round(self.word_count / 250)
### END SOLUTION
```

<!-- grade_id: cell-dcfdee410c3d03df, points: 1, grade: true, solution: false -->
```python
n1 = Novel("1984", "George Orwell", 1949, "dystopian", 89000)

assert n1.genre == "dystopian"
assert n1.word_count == 89000
assert n1.title == "1984"

rt1 = n1.get_reading_time()
assert rt1 == 356
```

<!-- grade_id: cell-928ef6f0e3f9504f, points: 1, grade: true, solution: false -->
```python
n2 = Novel("The Bluest Eye", "Toni Morrison", 1970, "fiction", 52000)

assert n2.genre == "fiction"
assert n2.word_count == 52000
assert n2.title == "The Bluest Eye"

rt2 = n2.get_reading_time()
assert rt2 == 208
```

<!-- grade_id: cell-493047d56f777541, points: 1, grade: true, solution: false -->
```python
# Check that it's the right type
assert isinstance(n1, Novel)
assert isinstance(n2, Novel)
# Check that it's inherited correctly
assert isinstance(n1, Text)
assert isinstance(n2, Text)
```

### Q3: Create a `Library` class

Now, create a class called `Library`, which has the following attributes:

- `location`: a `str`
- `books`: a `list` of `Novel` instances. 

This class should also have two methods:

- `organize_collection_by_year`: creates a new attribute (`self.collection_sorted_by_year`) that contains a *dictionary*, where each *key* is a *year*, and each *value* is a `list` of `Novel` instances published in that year.
- `organize_collection_by_genre`: creates a new attribute (`self.collection_sorted_by_genre`) that contains a *dictionary*, where each *key* is a *genre*, and each *value* is a `list` of `Novel` instances published in that year.

<!-- grade_id: cell-492a120752cffb3a, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
class Library:
    def __init__(self, location, books):
        self.location = location
        self.books = books
        self.collection_sorted_by_year = {}
        self.collection_sorted_by_genre = {}

    def organize_collection_by_year(self):
        for book in self.books:
            if book.year in self.collection_sorted_by_year:
                self.collection_sorted_by_year[book.year].append(book)
            else:
                self.collection_sorted_by_year[book.year] = [book]

    def organize_collection_by_genre(self):
        for book in self.books:
            if book.genre in self.collection_sorted_by_genre:
                self.collection_sorted_by_genre[book.genre].append(book)
            else:
                self.collection_sorted_by_genre[book.genre] = [book]
### END SOLUTION
```

<!-- grade_id: cell-9348d2cb890d091f, points: 1, grade: true, solution: false -->
```python
## set up novels
n1 = Novel("1984", "George Orwell", 1949, "dystopian", 89000)
n2 = Novel("The Bluest Eye", "Toni Morrison", 1970, "fiction", 52000)
n3 = Novel("Watership Down", "Richard Adams", 1970, "fantasy", 156000)
n4 = Novel("The Hobbit", "JRR Tolkien", 1937, "fantasy", 95000)

## create library
library = Library("San Diego", books = [n1, n2, n3, n4])

## check
assert library.location == "San Diego"
assert len(library.books) == 4
```

<!-- grade_id: cell-ecc396c391b020e4, points: 2, grade: true, solution: false -->
```python
# check for organizing by genre
library = Library("San Diego", books = [n1, n2, n3, n4])
library.organize_collection_by_genre()

assert len(library.collection_sorted_by_genre.keys()) == 3
assert "dystopian" in library.collection_sorted_by_genre
assert "fantasy" in library.collection_sorted_by_genre
assert "fiction" in library.collection_sorted_by_genre
assert len(library.collection_sorted_by_genre['fiction']) == 1
assert len(library.collection_sorted_by_genre['fantasy']) == 2

```

<!-- grade_id: cell-8ae550f47410f933, points: 2, grade: true, solution: false -->
```python
# check for organizing by year
library = Library("San Diego", books = [n1, n2, n3, n4])
library.organize_collection_by_year()

assert len(library.collection_sorted_by_year.keys()) == 3
assert 1949 in library.collection_sorted_by_year
assert 1970 in library.collection_sorted_by_year
assert 1937 in library.collection_sorted_by_year

assert len(library.collection_sorted_by_year[1949]) == 1
assert len(library.collection_sorted_by_year[1970]) == 2
```

## Submit!

Congratulations——you just implemented some `class`es in Python. Don't forget to *validate* and then *Submit* your assignment.